## Exercise 09.1 (checking data validity)

The Fibonacci series is valid only for $n \ge 0$. Add to the Fibonacci function in this notebook a check that raises an exception if $n < 0$. Try some invalid data cases to check that an exception is raised.

*Optional:* Use `pytest` to test that an exception *is* raised for some $n < 0$ cases.

<!-- grade_id: cell-982c3611466532a2, points: 0, grade: false, solution: true -->
```python
fib = [0 for i in range(100001)]
fib[1] = 1


def f(n):
    if n < 0:
        raise ValueError("not defined for negative numbers")
    if n >= len(fib):
        raise ValueError("maximum number of iterations is exceeded")

    if fib[n] != 0 or n == 0:
        return fib[n]
    fib[n] = f(n - 1) + f(n - 2)
    return fib[n]
```

<!-- grade_id: cell-9b97ea3e85d77f4a, points: 0, grade: true, solution: false -->
```python
## tests ##

# Perform some tests
assert f(0) == 0
assert f(1) == 1
assert f(2) == 1
assert f(3) == 2
assert f(10) == 55
assert f(15) == 610

# Check that ValueError is raised for n < 0
import pytest

with pytest.raises(ValueError):
    f(-1)
with pytest.raises(ValueError):
    f(-2)
```

## Exercise 09.2 (raising exceptions)

Modify your program from the bisection exercise in Activity 04 to raise an error if the maximum number of iterations is exceeded. Reduce the maximum allowed iterations to test that an exception is raised.

Add any other checks on the input data that you think are appropriate.

<!-- grade_id: cell-ca2c3ee62ad3b956, points: 0, grade: false, solution: true -->
```python
def my_f(x):
    """Evaluate polynomial function"""
    return x**5 / 10 + x**3 - 10 * x**2 + 4 * x + 7


def compute_root(f, x0, x1, tol, max_it):
    """Compute roots of a function using bisection"""
    if x0 > x1:
        raise ValueError("lower bound cannot be higher than upper bound")
    f0 = f(x0)
    f1 = f(x1)

    it = 0

    while it < max_it:
        x_mid = (x0 + x1) / 2
        f_mid = f(x_mid)

        if abs(f_mid) < tol:
            return x_mid, f_mid, it

        if f_mid * f0 < 0:
            x1 = x_mid
            f1 = f_mid
        else:
            x0 = x_mid
            f0 = f_mid

        it += 1

        if abs(x1 - x0) < tol:
            return x_mid, f_mid, it
    if it < max_it:
        return x_mid, f_mid, it
    else:
        raise RuntimeError("maximum number of iterations exceeded")
```

<!-- grade_id: cell-f6cf91ed5dd1736b, points: 0, grade: true, solution: false -->
```python
## tests ##

# Test with max_it = 30
x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=30)

# Test with max_it = 20
with pytest.raises(RuntimeError):
    x, f, num_it = compute_root(my_f, x0=0, x1=2, tol=1.0e-6, max_it=20)

# Test wih x0=2, x1=0
with pytest.raises(ValueError):
    x, f, num_it = compute_root(my_f, x0=2, x1=0, tol=1.0e-6, max_it=30)
```